# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 6

```python
%matplotlib inline
import numpy as np
import networkx as nx
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_count_equal
from numpy.testing import assert_almost_equal
import math
```

# First Course: chap 3.3 Friendship Paradox

1. Maak een netwerkx netwerk van Figuur 3.5.
2. Waarom is de kans dat je Tom vindt met de "random friend of picked person" methode $\frac{5}{21}$. Geef een heel helder antwoord.
3. Programmeer de functie `kans_op_knoop(G,knoop)` die voor elk netwerk, voor elke knoop de kans berekent dat je die knoop trekt met de "random friend of picked person" methode.
4. Schrijf een test die checked dat die kansen optellen tot 1.
5. **Advanced** In de een na laatste paragraaf van 3.3 op blz 76 wordt een andere manier om die kansen te berekenen gemeld, die dezelfde uitkomst geeft. Programmeer die ook, en vergelijk de uitkomsten.




2. Voor elke persoon in het netwerk is er een $\frac{1}{7}$ kans om willekeurig gekozen te worden. Er is bijvoorbeeld een $\frac{1}{7}$ kans dat John willekeurig gekozen wordt, en wanneer één van zijn vrienden willekeurig gekozen wordt, is er een $\frac{1}{2}$ kans dat dit Tom is. Deze twee kansen worden dus vermenigvuldigd. Voor Mary bestaat er geen kans dat een willekeurig gekozen vriend van haar Tom is, en voor Nancy geldt dat deze kans $\frac{1}{3}$ is. Voer dit uit voor alle nodes (waarvan één van hun vrienden Tom is, waar Tom zelf dus niet bij hoort) en je komt uit op een kans van $\frac{1}{7} \cdot \frac{1}{2} + \frac{1}{7} \cdot \frac{1}{3} + \frac{1}{7} \cdot \frac{1}{3} + \frac{1}{7} \cdot \frac{1}{2} = \frac{5}{21}$.

<!-- grade_id: FPa, points: 0, grade: false, solution: true -->
```python
Fig35=nx.Graph()
Fig35.add_nodes_from(['Mary', 'Nancy', 'Tom', 'Tara', 'John', 'Bob', 'Pam'])
Fig35.add_edges_from([('Mary', 'Nancy'), ('Nancy', 'Tom'), ('Nancy', 'Bob'), ('Tara', 'Bob'), ('Tom', 'John'), ('Tom', 'Bob'), ('Tom', 'Pam'), ('Pam', 'John')])

def kans_op_knoop(G,knoop):
    #WRITE YOUR CODE HERE
    random_knoop = 1 / G.number_of_nodes()
    neighbor_kans = 0
    for node in G:
        if knoop in G[node]:
            neighbor_kans += 1 / G.degree(node)
    return random_knoop * neighbor_kans
 

H= nx.karate_club_graph() 
# test
sum(kans_op_knoop(H,x) for x in H)

kans_op_knoop(Fig35, 'Tom')
```

Output:
```
0.23809523809523805
```

<!-- grade_id: FPt, points: 1, grade: true, solution: false -->
```python
assert isinstance(Fig35, nx.Graph)
H= nx.karate_club_graph()
assert 0 <= kans_op_knoop(H,1)<= 1

```

# First course

Maak de sommen 3.11-3.24.

11. 1. Tie between 2, 3, and 5, since they all have a degree of 3.
    2. Tie between 2 and 5.
    3. Node 2.
12. In this graph, all the nodes will need to connect to one singular node, the hub. This will mean that the avarage degree is $\frac{1 \cdot 9 + 9}{10} = 1.8$. The avarage square degree is $\frac{1^2 \cdot 9 + 9^2}{10} = 9$, meaning the heterogeinity parameter is $\frac{9}{1.8} = 5$.
13. 1. No heavy tailed distribution, there aren't a lot of people with an extraordinarily high shoe size.
    2. Heavy tailed distribution, there are quite a few people in the US that make more than the modal income.
    3. Heavy tailed distribution, thanks to people like celebrities. These nodes are still quite abundant and have a high degree.
    4. No heavy tailed distribution, since the majority of pair distances is low. A high distance is way less abundant.
14. I would not be.
15. 1. $1e08$ pages.
    2. $1e06$ pages.
    3. 10000 pages.

3.16

    What is the maximum degree and what does this mean?

    A: The degree in this study is the amount of sexual partners. The maximum degree is the maximum amount of sexual partners one reports in the survey of the study

    If you consider the subnetworks with nodes corresponding to males and females, respectively, do they have the same degree distribition? Why or why not?

    A: Males relatively report a higher degree of sexual partners than females. This could be explained by different social theories. One might be that it socially more acceptable in certain cultures for men to have more sexual partner than women, thus resulting in an overestimation of reports whereas women tend to lower their degree as a socially acceptable answer to the survey
    
20. The degree must be the same for every node. (regular network)

21. Targeted, since they will just remove the nodes with the highest degrees.  
22. D, since nodes with a high betweenness centrality lie on the shortest path of many nodes, deleting that node will cause the shortest paths of these nodes to be longer.
23. The one with the lowest clustering coëfficient, that way as many paths as possible will be disrupted.

```python
# Opdracht 17
Flights = nx.read_edgelist('./openflights_usa.edges')
nx.draw(Flights)

```

Output:
```
File <class 'networkx.utils.decorators.argmap'> compilation 5:3, in argmap_read_edgelist_1(path, comments, delimiter, create_using, nodetype, data, edgetype, encoding, backend, **backend_kwargs)
      1 import bz2
      2 import collections
----> 3 import gzip
      4 import inspect
      5 import itertools

File ~/anaconda3/envs/Teaching2024/lib/python3.10/site-packages/networkx/utils/decorators.py:193, in open_file.<locals>._open_file(path)
    189 else:
    190     # could be None, or a file handle, in which case the algorithm will deal with it
    191     return path, lambda: None
--> 193 fobj = _dispatch_dict[ext](path, mode=mode)
    194 return fobj, lambda: fobj.close()

FileNotFoundError: [Errno 2] No such file or directory: './openflights_usa.edges'
```

<!-- grade_id: fc, points: 14, grade: true, solution: true -->
```python
def Vraag1(G):
    total = sum([G.degree(node) for node in G])
    return total / G.number_of_nodes()

print(f"Average number of routes: {Vraag1(Flights)}")

def Vraag2(G):
    num_routes = list(G.degree())
    num_routes.sort(key=lambda a: a[1], reverse = True)
    return ", ".join([info[0] for info in num_routes[:5]])

print(f"Top 5 airports with highest routes: {Vraag2(Flights)}.")

def Vraag3(G):
    one_route_ports = [port for port in G.degree() if port[1] == 1]
    return len(one_route_ports)

print(f"There are {Vraag3(Flights)} airports that serve only a single route.")

def Vraag4(G):
    return max(G.nodes, key=nx.closeness_centrality(G).get)

print(f"The airport with the highest closeness centrality is {Vraag4(Flights)}")

def Vraag5(G):
    betweenness = nx.betweenness_centrality(G)
    betweenness = dict(sorted(betweenness.items(), key=lambda item: item[1], reverse = True))
    return list(betweenness)[0]

print(f"The airport with the highest betweenness centrality is {Vraag5(Flights)}")

def Vraag6(G):
    avg_sq_degree = sum([(G.degree(node)) ** 2 for node in G]) / G.number_of_nodes()
    avg_degree = Vraag1(G)
    return avg_sq_degree / (avg_degree ** 2)

print(f"The heterogeneity of the network is equal to {Vraag6(Flights)}.")
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[5], line 5
      2     total = sum([G.degree(node) for node in G])
      3     return total / G.number_of_nodes()
----> 5 print(f"Average number of routes: {Vraag1(Flights)}")
      7 def Vraag2(G):
      8     num_routes = list(G.degree())

NameError: name 'Flights' is not defined
```

<!-- grade_id: fc324, points: 6, grade: true, solution: true -->
```python
#WRITE YOUR CODE HERE
# 3.18

W = nx.read_graphml("./enwiki_math.graphml.gz")


# 1
print(sum(W.out_degree(node) for node in W) / len(W.nodes()))
print(sum(W.in_degree(node) for node in W) / len(W.nodes()))

# 2

max_k_in_node = max(W.nodes, key=W.in_degree(W.nodes))
print(W.nodes[max_k_in_node]['title'])
# 3

max_k_in_node = max(W.nodes, key=W.out_degree(W.nodes))
print(W.nodes[max_k_in_node]['title'])

# 4

# max out degree, no because for airports it's not uncommon that one airport would receive a lot of planes, like Newark, but would not send out much.


# 5
# 6


```

Output:
```
File ~/anaconda3/envs/Teaching2024/lib/python3.10/gzip.py:58, in open(filename, mode, compresslevel, encoding, errors, newline)
     56 gz_mode = mode.replace("t", "")
     57 if isinstance(filename, (str, bytes, os.PathLike)):
---> 58     binary_file = GzipFile(filename, gz_mode, compresslevel)
     59 elif hasattr(filename, "read") or hasattr(filename, "write"):
     60     binary_file = GzipFile(None, gz_mode, compresslevel, filename)

File ~/anaconda3/envs/Teaching2024/lib/python3.10/gzip.py:174, in GzipFile.__init__(self, filename, mode, compresslevel, fileobj, mtime)
    172     mode += 'b'
    173 if fileobj is None:
--> 174     fileobj = self.myfileobj = builtins.open(filename, mode or 'rb')
    175 if filename is None:
    176     filename = getattr(fileobj, 'name', '')

FileNotFoundError: [Errno 2] No such file or directory: './enwiki_math.graphml.gz'
```

<!-- grade_id: 3_19, points: 0, grade: false, solution: true -->
```python
# SOM 3.19
from statistics import mean

def  avg_degree_neighbors(G, friend):
    return mean([G.degree(node) for node in G.neighbors(friend)])
    
H= nx.karate_club_graph() 

{x: avg_degree_neighbors(H,x) for x in H}   

# paradox op de karateclub
(sum(avg_degree_neighbors(H,x) for x in H )/len(H), 
 sum(H.degree(x) for x in H)/len(H), 
2* H.number_of_edges()/H.number_of_nodes())

```

Output:
```
(9.6102100153787, 4.588235294117647, 4.588235294117647)
```