# Assignment Netwerkanalyse

## Notebook made by

**Gebruik graag dit formaat**

* Voor de namen:  voornaam rest van je naam, voornaam rest van je naam,....
* je studentnummers: hetzelfde: scheidt met `,`
* je emails: hetzelfde: scheidt met `,`
* voor je groep: alleen de hoofdletter (dus voor de groep van Marx zou je `A` kiezen)

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 1

# Deel 1: Programmeren

Voor deze opdrachten maken we gebruik van Networkx, de documentatie daarvan is te vinden via [deze link](https://networkx.github.io/documentation/stable/). Kijk vooral naar welke functies zij bieden, zodat je niet dingen doet die zij al voor je doen (tenzij dat natuurlijk de opdracht is).

Staat er `ModuleNotFoundError: No module named X`, doe dan in een cell `!pip install X` (natuurlijk is X hier iets van `matplotlib`, `networkx`, `scipy` of `nose`) en run de cell dan opnieuw.

Mocht dat niet werken, probeer dan `!pip3 install X` (werkt alleen bij Mac) of restart de kernel en probeer opnieuw! Mocht dat allemaal ook niet werken, roep dan even een TA!

```python
# Installeer benodigde packages
%matplotlib inline
import networkx as nx
import matplotlib.pyplot as plt
import numpy as np
import scipy
from nose.tools import assert_equal
from numpy.testing import assert_almost_equal, assert_array_almost_equal
```

## Vraag 1
<p>Maak een networkx ongerichte graaf `G` met de volgende knopen en lijnen:</p>
<center>
    V = {1, 2, 3, 4, 5} en 
    E = { (1, 2) , (1, 4) , (2, 3) , (2, 4) , (2, 5) , (3, 4) , (3, 5) }
</center>
De uitkomst van de volgende cel is een mooi plaatje van een netwerkje met... ja echt 5 knopen.


<!-- grade_id: Vraag1, points: 0, grade: false, solution: true -->
```python
G = nx.Graph()

V = [1, 2, 3, 4, 5]
G.add_nodes_from(V)

E = [(1,2), (1,4), (2,3), (2,4), (2,5), (3,4), (3,5)]
G.add_edges_from(E)

nx.draw(G, with_labels=True)
```

<!-- grade_id: Vraag1_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(nx.is_directed(G), False)
```

## Vraag 2
Maak op basis van het netwerk `G` een dictionary genaamd `neighborDict` met als key een node en als value de verzameling  met alle buren van die knoop.
Gebruik dict comprehensie.

Bijvoorbeeld: `neighborDict[5] = {2, 3}`

<!-- grade_id: Vraag2, points: 0, grade: false, solution: true -->
```python
#deze dict comprehension loopt door de nodes in het netwerk G
#Elke node wordt een key in de dictionary
#de value is een lijst die ontstaat door te loopen over alle buren van deze node
#de lijst wordt vervolgens omgezet in een set

neighborDict = {node: {x for x in G.neighbors(node)} for node in G.nodes()}

neighborDict[5]
```

<!-- grade_id: Vraag2_antwoord, points: 1, grade: true, solution: false -->
```python
# Onderstaande tests helpen je bij het vinden van het juiste antwoord
# Als hier iets mis gaat, probeer dan de foutmelding te begrijpen en jouw antwoord aan te passen
# Ook op het tentamen geven we zulke hints
# De echte check of je antwoord correct is is geheim. Die staat hieronder maar zie jij niet.


assert_equal(type(neighborDict), dict) # we testen hier of jouw antwoord van het juiste type is.
assert_equal(len(neighborDict),5)  # we testen hier of je elke knoop in je dict hebt gestopt en niks meer
assert_equal(set(neighborDict.keys()), set(G.nodes)) # en hier doen we dat nog een keer
for knoop in neighborDict:  # hier testen we of elke verzameling buren van elke knoop in jouw antword echt knopen uit G zijn.
    assert neighborDict[knoop].issubset(set(G.nodes))


```

## Vraag 3
Maak een functie `make_complete(G)` die teruggeeft hoeveel lijnen toegevoegd moeten worden aan een willekeurig  ongericht netwerk `G` (die misschien al lijnen bevat) om er een volledige graaf van te maken. Zorg dat het resultaat een integer is.


<!-- grade_id: Vraag3, points: 0, grade: false, solution: true -->
```python
def make_complete(G):
    
    #het maximale aantal edges in een graph is het aantal nodes - 1
    max_degree = G.number_of_nodes() - 1 
    
    #een compleet netwerk wordt berekend aan de hand van de formule: (aantal nodes * maximale degree) / 2
    complete_network = (G.number_of_nodes() * max_degree) / 2 
    
    #het antwoord is de hoeveelheid edges van het complete netwerk - de hoeveelheid edges die het netwerk al heeft
    return int(complete_network - G.number_of_edges())

# test 1
print(make_complete(G))
# test 2
ER = nx.gnm_random_graph(100, 200) # Zie https://networkx.github.io/documentation/stable/auto_examples/graph/plot_erdos_renyi.html
nx.draw(ER, with_labels=True);
make_complete(ER)
```

<!-- grade_id: Vraag3_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(make_complete(G)), int) # jouw antwoord is een geheel getal, een integer
```

## Vraag 4
Maak een functie `my_degree(G)` die een dictionary maakt met als sleutels alle knopen van `G`    en als value de degree van de sleutel in `G`. Het resultaat is hetzelfde als `dict(G.degree())`.

* Gebruik een dict comprehensie!
* Je mag alleen de networkx functie `G.has_edge()` gebruiken. 
* Dit kan gewoon in 1 regel. Ga geen gekke loops of zo gebruiken. Wees elegant, en vermijd zo foutjes. Denk liever zelf eerst na vordat je iemand nadoet die jij niet kent en die iets op het internet heeft geroepen.

Zie: https://en.wikipedia.org/wiki/Degree_(graph_theory)

Voorbeeld resultaat op ons eigen netwerk `G`: `{1: 2, 2: 4,...}`

<!-- grade_id: Vraag4, points: 0, grade: false, solution: true -->
```python
def my_degree(G):
    
    return {node: G.degree(node) for node in G.nodes()}

# test 1
print('G:', my_degree(G))

# test 2
ER = nx.gnm_random_graph(10, 20) # Zie https://networkx.github.io/documentation/stable/auto_examples/graph/plot_erdos_renyi.html
nx.draw(ER, with_labels=True);
my_degree(ER)
```

<!-- grade_id: Vraag4_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(my_degree(G), dict(G.degree()))
```

## Vraag 5
Teken de subgraven van `G` geinduceerd door {1,2,3} en {2,5,3} en {5,4} en geef deze graven de namen respectievelijk `I`, `J` en `K`

<!-- grade_id: Vraag5, points: 0, grade: false, solution: true -->
```python
I = G.subgraph([1,2,3])
J = G.subgraph([2,5,3])
K = G.subgraph([5,4])

#WRITE YOUR CODE HERE
plt.figure(1)
nx.draw(I)
plt.figure(2)
nx.draw(J)
plt.figure(3)
nx.draw(K)
plt.show()
```

<!-- grade_id: Vraag5_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(sorted(list(I.nodes())), [1, 2, 3])
assert_equal(sorted(list(J.nodes())), [2, 3, 5])
assert_equal(sorted(list(K.nodes())), [4, 5])
```

## Vraag 6
Maak de functie `adj_matrix(G)`  en `adj_matrix2(G)` die allebei een adjacency matrix maken. In het geval van de graaf `G` die je hebt gemaakt bij vraag 1, heeft de matrix van deze  vorm (alleen dan in de vorm van een lijst van lijsten):
    <pre>
        1 2 3 4 5
      1 0 1 0 1 0
      2
      3
      4
      5
    </pre>
    
Een aantal eisen:
* Zorg dat de functies werken voor elk netwerk, niet alleen `G`. (We kijken de functies ook na met andere netwerken dan `G`.
* Gebruik enkel list comprehensies.
* Je mag in beide functies niet de networkx functie `nx.adjacency_matrix` gebruiken.
* Bij `adj_matrix(G)` moet je de networkx functie `G.has_edge` gebruiken en  geen enkele andere networkx functie.
* Bij `adj_matrix2(G)` moet je de networkx functie `G.neighbors` (dit is niet hetzelfde als `nx.neighbors`!) gebruiken en  geen enkele andere networkx functie.
* Zorg dat de nodes altijd gesorteerd zijn, zodat de matrix op volgorde is
* Hint:`int(True)==1` en `int(False)==0`, en dus ook `int(5==4)==0`.
    
Je antwoord voor netwerk `G` zal er zo uit zien:

```
[[0, 1, 0, 1, 0],
 [....],
       ....
 [....]
] 
```

<!-- grade_id: Vraag6, points: 0, grade: false, solution: true -->
```python
def adj_matrix(G):
    
    matrix = []
    
    for row_nodes in G:
        rows = []
        for column_nodes in G:
            if G.has_edge(row_nodes, column_nodes):
                rows.append(1)
            else:
                rows.append(0)
        matrix.append(rows)
    return matrix

def adj_matrix2(G):
    
    matrix = []
    
    for row_nodes in G:
        rows = []
        for column_nodes in G:
            if column_nodes in G.neighbors(row_nodes):
                rows.append(1)
            else:
                rows.append(0)
        matrix.append(rows)
    return matrix

adj_matrix(G),adj_matrix2(G), adj_matrix(G)==adj_matrix2(G)

# test het zelf ook met de ER graaf.
```

<!-- grade_id: Vraag6_test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(adj_matrix(G)), list)
assert_equal(type(adj_matrix(G)[0]), list)
assert_equal(type(adj_matrix(G)[0][0]), int)
assert_equal(type(adj_matrix2(G)), list)
assert_equal(type(adj_matrix2(G)[0]), list)
assert_equal(type(adj_matrix2(G)[0][0]), int)
assert_equal(type(adj_matrix2(G)[0][0]), int)
assert_array_almost_equal(adj_matrix(G), adj_matrix2(G))

```

# Opgaven uit First Course

1. Alles uit Hfd 0
2. Hfd 1: alles tot en met 21.
3. Maak opgave 21 in `networkx`
        * maak het netwerk, met `G.add_edges`
        * print het
        * maak ook de adjecency matrix als een lijst van lijsten
        * check dat deze adjencency matrix hetzelfde is als een adjececy matrix die networkx maakt.
4. Schrijf nu ook checks voor alle vragen 1.17-1.20 in Python en controleer je antwoord.

<!-- grade_id: FCa, points: 1, grade: true, solution: true -->
**answers chapter 0**

0.1	C

0.2	A

0.3	A flights that departed from on airport to the other airport

0.4	D

0.5	De link between a ‘friend’ connection between two users, this is undirected.

0.6	A user who decides to follow another users without that user following him/her back, this is directed.


**answers chapter 1**

1.2 a link can connect a maximum of two nodes, a node can have a maximum of N-1 links

1.3  -

1.4 the in-degree and out-degree would both be 2

1.5 –

1.6 D

1.7 D

1.8 B

1.9 A

1.10 D

1.11 a complete network of 300 links

1.12 N(N-1) / 2

1.13 N1 x N2

1.14 A

1.15 since there can’t be an link between a node itself, the network must be complete in the case of N(N-1)/2 links. 

1.16 D

1.17 six nodes, 10 edges, and no self-loops

1.18 node F has outgoing links to every other node, there are no nodes with in-links from every node.

1.19 node C

1.20 in-strength is 5 and out-strength is 0


# FC vraag 1.13

* Maak deze vraag ook in networkx.
    * Maak een functie `Q113(N1,N2)` die een network teruggeeft zoals gevraagd
    * Met welke check check je dat het "volledig" (voor een bipartiet netwerk natuurlijk) is?
    
* **Hint** Google eens wat.

<!-- grade_id: fc113a, points: 0, grade: false, solution: true -->
```python
def Q113(N1,N2):
    '''Return a maximal bipartite network with N1 nodes on one side and N2 on the other.'''
    
    return nx.complete_bipartite_graph(N1, N2)
    
def maximaal_bipartiet(BG):
    '''Return True if BG is a maximal bipartite network. False in all other cases.'''
    
    nodes = nx.bipartite.sets(BG) #will retrieve the two sets of nodes from the bipartite network
    N1, N2 = nodes #stores the two sets of nodes in seperate variables
    
    for node1 in N1:
        for node2 in N2:
            if not BG.has_edge(node1, node2): #checks if there's a missing edge in the complete graph
                return False 
    return True #returns True if graph is complete
    
#test
maximaal_bipartiet(Q113(5,10))
```

<!-- grade_id: fc113t, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q113(5,10),nx.Graph)
assert isinstance(maximaal_bipartiet(Q113(5,10)), bool)


    
 
```

<!-- grade_id: adjmatrixchecksd, points: 1, grade: true, solution: true -->
```python
#code opgave 21

WG = nx.DiGraph()
edges_with_weights = [('A', 'B', 1), ('B', 'C', 2), ('D', 'B', 1), ('D', 'E', 1), ('E', 'F', 1), 
                      ('F', 'A', 2), ('F', 'B', 1), ('F', 'C', 3), ('F', 'D', 1), ('F', 'E', 1)]
WG.add_weighted_edges_from(edges_with_weights)
nx.draw(WG, with_labels = True)
```

```python
# adjacency matrix als lijst van lijsten voor graph opgave 21
matrix = []

for nodes_row in WG.nodes():
    rows = []
    for nodes_column in WG.nodes():
        edge = (nodes_row, nodes_column)
        if edge in WG.edges(): #controleert of the link in het netwerk zit
            rows.append(WG.get_edge_data(nodes_row, nodes_column)['weight']) #als de links in het netwerk zit, wordt de bijbehorende weight toegevoegd
        else:
            rows.append(0) #als er geen link is, wordt 0 toegevoegd
    matrix.append(rows)

print(matrix)
            
#WG.get_edge_data('F', 'C')['weight']
```

<!-- grade_id: FCa21, points: 1, grade: true, solution: true -->
```python
#check functie voor 1.17

WG_matrix = [[0, 1, 0, 0, 0, 0], [0, 0, 2, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 1, 0, 0, 1, 0], [0, 0, 0, 0, 0, 1], [2, 1, 3, 1, 1, 0]]

def network_structure(adj_matrix):
    
    number_nodes = len(adj_matrix)
    number_edges = 0
    self_loops = 0
    
    for rows in adj_matrix:
        for values in rows:
            if values >= 1:
                number_edges += 1
                
    for index in range(len(adj_matrix)):
        row = adj_matrix[index]
        if row[index] != 0:
            self_loops += 1

    print(f'This adjacency matrix contains {number_nodes} nodes, {number_edges} edges, and {self_loops} self-loops')

#check functie voor 1.18

def di_network_structure(WG):
    only_incoming = False
    only_incoming_nodes = []
    for node in WG.nodes():
        for pairs in WG.edges():
            if (node, node) not in pairs and WG.in_degree(node) == len(WG.nodes()) - 1:
                only_incoming = True
                only_incoming_nodes.append(node)
            
    only_outgoing = False
    only_outgoing_nodes = []
    for node in WG.nodes():
        for pairs in WG.edges():
            if (node, node) not in pairs and WG.out_degree(node) == len(WG.nodes()) - 1:
                only_outgoing = True
                if node not in only_outgoing_nodes:
                    only_outgoing_nodes.append(node)
    
    if only_incoming == True:
        print(f'This graph has the following nodes with incoming links from every other node: {only_incoming_nodes}')
    else:
        print('This graph has no nodes with only incoming links from every other node')

    if only_outgoing == True:
        print(f'This graph has the following nodes with outgoing links to every other node: {only_outgoing_nodes}')
    else:
        print('This graph has no nodes with outgoing links to every other node')
        
#check functie voor 1.19

def sink(WG):
    sinks = []
    sinks_present = False

    for node in WG.nodes():
        if WG.in_degree(node) != 0 and WG.out_degree(node) == 0:
            sinks_present = True
            sinks.append(node)
            
    if sinks_present == True:
        print(f'This graph has the following sinks: {sinks}')
    else:
        print('This graphs has no sinks')
        
#check functie voor 1.20

def strength(node, WG):
    
    in_strength = WG.in_degree(node, weight = 'weight')
    out_strength = WG.out_degree(node, weight = 'weight')
    
    print(f'node {node} heeft een in-strength van: {in_strength} en een out-strength van: {out_strength}')


network_structure(WG_matrix)
di_network_structure(WG)
sink(WG)
strength('C', WG)
```

# Verzamelingen

* Dubbel klik eens op deze cel zodat je de code ziet.
* Leer dan de latex codes voor de mooie symbolen, en zorg dat je zelf nu ook netjes antwoord kunt geven.
* Formules zet je altijd tussen dollartekens.
    * Als je het in het midden van de regel wil tussen dubbele dolalartekens
* Elk commando begint altijd met `\`. Spaties maken niks uit. Latex regelt hoe het er netjes uit komt te zien.

### Opgaven

1. $\{1,2,3\}\cap\{2,3,4\}=\{2,3\}$
1. $\{1,2,3\}\cup\{2,3,4\}= \{1,2,3,4\}$
1. $\{1,2,3\}\setminus\{2,3,4\}= \{1\}$
1. $\emptyset\setminus\{2,3,4\}=\{\}$
2. $\{1,2,3\}\times\{2,3,4\}=\{\{1,2\}, \{1,3\}, \{1,4\}, \{2,2\}, \{2,3\}, \{2,4\}, \{3,2\}, \{3,3\}, \{3,4\}\}$
3. De machtsverzameling van $\{1,2,3,4\}$ is ....\
$\{\{\}, \{1\}, \{2\}, \{3\}, \{4\}, \{1,2\}, \{1,3\}, \{1,4\}, \{2,3\}, \{2,4\}, \{3,4\}, \{1,2,3\}, \{1,2,4\}, \{2,3,4\}, \{1,3,4\}, \{1,2,3,4\}\}$
     * Check dat je echt $2^4$ veel deelverzamelingen hebt gevonden.
     * het heet niet voor niks *machts* verzameling ;-)

#### Teken ook netjes de Venn diagrammen
#### Check je antwoorden ook in Python
#### Maak zelf opgaven en leg die aan je groepsleden voor

# Comprehensies

* Zorg dat je list, set en dict comprehensies echt goed in je vingers krijgt!
* Het is heerlijk als je het kunt.

### Goede uitleg

Ik [Googlede op python comprehension tutorial](https://www.google.com/search?q=python+comprehension+tutorial) en de [video](https://www.youtube.com/watch?v=3dt4OGnU5sM) is prima, maar ook de [datacamp tutorial](https://www.datacamp.com/community/tutorials/python-list-comprehension) of de [real python](https://realpython.com/list-comprehension-python/) tutorial.

### Wat opgaven

Je gaat allemaal comprehensies doen op een gedicht van annie MG Schmidt. Het staat in de cel hieronder als string in de variabele `AMGS` en daarna als lijst van woorden in `t`.

In je tests gebruik je allemaal string functies, en ook `in`. 

#### list comprehensies

1. de lijst van alle woorden in `AMGS` beginnend met een `i`. (antwoord beneden)
2. zelfde maar dan beginnend met `i` of `I`. (antwoord beneden)
3. de lijst van alle woorden in `AMGS` van lengte meer dan 3
4. de lijst van alle woorden in `AMGS` waar een `a` in zit
5. nu waar juist geen `a` in zit
6. nu waar als er een `a` in zit ook een `b` in zit (wat betekent dit precies he?)
7. alle woorden zonder klinker.
8. Voor elk woord zijn lengte.
9. De lijst met het hele gedicht maar nu waar elk woord met een hoofdletter begint.
10. De lijst die net zo lang is als het gedicht (dus `t`), maar dan op elke positie of het woord op die positie gelijk is aan `ik`. Dus de waardes zijn alleen maar `True` en `False`.
    * Gebruik `sum` om te berekenen hoe vaak het woord `ik` voorkomt.
11. Verzien er zelf 10 bij, en maak die ook, en geef ze aan je groepsgenoeten.

#### set comprehensies

1. De verzameling van alle unieke woorden in `AMGS`.
2. De verzameling van alle unieke woorden in `AMGS` met `ch` er in.
3. De verzameling van alle unieke woorden in `AMGS`, maar dan geschreven in hoofdletters.
4. Alle woordlengtes die in `AMGS` voorkomen.


### Dict comprehensies

1. De dict met voor elk uniek "woord" in `AMGS` zijn lengte. Dus het woord is de key en de lengte de waarde. Let op , woorden zijn dus de elementen in `t`. (antwoord beneden)
2. De dict met voor elk uniek "woord" in `AMGS` de verzameling karakters in dat woord.
3. De dict met voor elk uniek "woord" in `AMGS`  de "schoongemaakte" versie van dat woord (zonder leestekens, etc, gebruik een lange lijst als `a.replace(':','').replace(',','')`. Zet het woord ook in klein eletters.

    
### Nesting

3. De dict met voor elk uniek woord, de verzameling klinkers in dat  woord. 
    *   (antwoord beneden, probeer eerst minstens een half uur zelf)
4.  De dict met voor elke positie in `AMGS`, de verzameling met het woord voor die positie  en het woord erna.
    (antwoord ook beneden, niet meteen gan kijken.)
5. DE dict met voor elk woord de verzameling posities waar het voorkomt. Tricky maar leuk!   Hint loop over de posities. 
6. Na de vorige is deze heel makkelijk: DE dict met voor elk woord  hoe vaak het voorkomt. 


```python
AMGS='''
Ik wil niet meer, ik wil niet meer!
Ik wil geen handjes geven!
Ik wil niet zeggen elke keer:
Jawel mevrouw, jawel meneer...
nee, nooit meer in m'n leven!
Ik hou m'n handen op m'n rug
en ik zeg lekker niks terug!

Ik wil geen vieze havermout,
ik wil geen tandjes poetsen!
Ik wil lekker knoeien met het zout,
ik wil niet aardig zijn, maar stout
en van de leuning roetsen
en schipbreuk spelen in de teil
en ik wil spugen op het zeil!

En heel hard stampen in een plas
en dan m'n tong uitsteken
en morsen op m'n nieuwe jas
en ik wil overmorgen pas
weer met twee woorden spreken!
En ik wil alles wat niet mag,
de hele dag, de hele dag!

En ik wil op de kanapee
met hele vuile schoenen
en ik wil aldoor gillen: nee!
En ik wil met de melkboer mee
en dan het paardje zoenen.
En dat is alles wat ik wil
en als ze kwaad zijn, zeg ik: Bil!'''

t=AMGS.split()
len(t),t[:10]
```

```python
# probeer eens alleen t[0]. en dan shift return, en probeer eens wat van de suggesties
# om erachter te komen wat iets doet typ je t[0].isupper? en dan shift enter.
# Ideaal he! Dat kan bij elk object. t[0] is natuurlijk een string, dus je krijgt string functies te zien. 
t[0].isupper()
```

```python
# "in" is ook heel handig in Python

print(t[1])
'i' in t[1], 'W' in t[1], 'W' in t[1].upper()
```

```python
#de lijst van alle woorden in `AMGS` beginnend met een `i`. (antwoord beneden)
[w for w in t if w.startswith('i')]
```

```python
#2. zelfde maar dan beginnend met `i` of `I`. (antwoord beneden)
# Verzin minstens 2 andere manieren om exact hetzelfde te krijgen
[w for w in t if w.startswith('i') or w.startswith('I')]
```

```python
#  De dict met voor elk uniek woord in `AMGS` zijn lengte.

{w:len(w) for w in t}
```

```python
# De dict met voor elk uniek woord, de verzameling klinkers in dat  woord.

{w:{l for l in w if l in 'aeiou'} 
    for w in set(t)} # verbeter hem zelf zodat hoofdletters ook mee worden genomen

```

```python
#  De dict met voor elke positie in `AMGS`, de verzameling met het woord voor die positie  en het woord erna.

{i:{t[i-1],t[i+1]} for i in range(1,len(t)-1)}
```

<!-- grade_id: TRICKY, points: 0, grade: false, solution: true -->
```python
#  DE dict met voor elk woord de verzameling posities waar het voorkomt. Tricky maar leuk!    
l = {w : {i for i in range(1, len(t)-1) if t[i] == w} for w in t}
print(l)
```

```python

```