# Assignment Netwerkanalyse

## Notebook made by

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem
__Groep__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd.

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

# Week 1

# Deel 1: Programmeren

Voor deze opdrachten maken we gebruik van Networkx, de documentatie daarvan is te vinden via [deze link](https://networkx.github.io/documentation/stable/). Kijk vooral naar welke functies zij bieden, zodat je niet dingen doet die zij al voor je doen (tenzij dat natuurlijk de opdracht is).

Staat er `ModuleNotFoundError: No module named X`, doe dan in een cell `!pip install X` (natuurlijk is X hier iets van `matplotlib`, `networkx`, `scipy` of `nose`) en run de cell dan opnieuw.

Mocht dat niet werken, probeer dan `!pip3 install X` (werkt alleen bij Mac) of restart de kernel en probeer opnieuw! Mocht dat allemaal ook niet werken, roep dan even een TA!

```python
# Installeer benodigde packages
!pip install nose matplotlib scipy nose
%matplotlib inline
import networkx as nx
import matplotlib.pyplot as plt
import numpy as np
import scipy
from nose.tools import assert_equal
from numpy.testing import assert_almost_equal, assert_array_almost_equal
```

## Vraag 1
<p>Maak een networkx ongerichte graaf `G` met de volgende knopen en lijnen:</p>
<center>
    V = {1, 2, 3, 4, 5} en
    E = { (1, 2) , (1, 4) , (2, 3) , (2, 4) , (2, 5) , (3, 4) , (3, 5) }
</center>
De uitkomst van de volgende cel is een mooi plaatje van een netwerkje met... ja echt 5 knopen.


<!-- grade_id: Vraag1, points: 0, grade: false, solution: true -->
```python
G = None

G = nx.Graph()
G.add_nodes_from([1, 2, 3, 4, 5])
G.add_edges_from([
    (1, 2),
    (1, 4),
    (2, 3),
    (2, 4),
    (2, 5),
    (3, 4),
    (3, 5),
])

nx.draw(G, with_labels=True)
```

<!-- grade_id: Vraag1_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(nx.is_directed(G), False)
```

## Vraag 2
Maak op basis van het netwerk `G` een dictionary genaamd `neighborDict` met als key een node en als value de verzameling  met alle buren van die knoop.
Gebruik dict comprehensie.

Bijvoorbeeld: `neighborDict[5] = {2, 3}`

<!-- grade_id: Vraag2, points: 0, grade: false, solution: true -->
```python
neighborDict = None
#WRITE YOUR CODE HERE
neighborDict = {node: set(G.neighbors(node)) for node in G.nodes()}
neighborDict
```

<!-- grade_id: Vraag2_antwoord, points: 1, grade: true, solution: false -->
```python
# Onderstaande tests helpen je bij het vinden van het juiste antwoord
# Als hier iets mis gaat, probeer dan de foutmelding te begrijpen en jouw antwoord aan te passen
# Ook op het tentamen geven we zulke hints
# De echte check of je antwoord correct is is geheim. Die staat hieronder maar zie jij niet.


assert_equal(type(neighborDict), dict) # we testen hier of jouw antwoord van het juiste type is.
assert_equal(len(neighborDict),5)  # we testen hier of je elke knoop in je dict hebt gestopt en niks meer
assert_equal(set(neighborDict.keys()), set(G.nodes)) # en hier doen we dat nog een keer
for knoop in neighborDict:  # hier testen we of elke verzameling buren van elke knoop in jouw antword echt knopen uit G zijn.
    assert neighborDict[knoop].issubset(set(G.nodes))


```

## Vraag 3
Maak een functie `make_complete(G)` die teruggeeft hoeveel lijnen toegevoegd moeten worden aan een willekeurig  ongericht netwerk `G` (die misschien al lijnen bevat) om er een volledige graaf van te maken. Zorg dat het resultaat een integer is.


<!-- grade_id: Vraag3, points: 0, grade: false, solution: true -->
```python
def make_complete(G):
    "Geeft terug hoeveel lijnen toegevoegd moeten worden aan een willeukeurig ongericht netwerk om een volledige graaf te maken"
    # aantal lijnen in de grafiek
    aantal_bestaande_lijnen = G.number_of_edges()

    # pak aantal punten/knopen in de grafiek
    aantal_bestaande_knopen = G.number_of_nodes()

    # totaal aantal lijnen mogelijk
    # antwoord is het verschil tussen de 2
    P = aantal_bestaande_knopen
    totaal_lijnen_mogelijk = (P * (P - 1) / 2)

    lijnen_nog_nodig = totaal_lijnen_mogelijk - aantal_bestaande_lijnen
    return int(lijnen_nog_nodig)




# test 1
print(make_complete(G))
# test 2
ER = nx.gnm_random_graph(100, 200) # Zie https://networkx.github.io/documentation/stable/auto_examples/graph/plot_erdos_renyi.html
nx.draw(ER, with_labels=True);
make_complete(ER)
```

<!-- grade_id: Vraag3_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(type(make_complete(G)), int) # jouw antwoord is een geheel getal, een integer
```

## Vraag 4
Maak een functie `my_degree(G)` die een dictionary maakt met als sleutels alle knopen van `G`    en als value de degree van de sleutel in `G`. Het resultaat is hetzelfde als `dict(G.degree())`.

* Gebruik een dict comprehensie!
* Je mag alleen de networkx functie `G.has_edge()` gebruiken.
* Dit kan gewoon in 1 regel. Ga geen gekke loops of zo gebruiken. Wees elegant, en vermijd zo foutjes. Denk liever zelf eerst na vordat je iemand nadoet die jij niet kent en die iets op het internet heeft geroepen.

Zie: https://en.wikipedia.org/wiki/Degree_(graph_theory)

Voorbeeld resultaat op ons eigen netwerk `G`: `{1: 2, 2: 4,...}`

<!-- grade_id: Vraag4, points: 0, grade: false, solution: true -->
```python
def my_degree(G):
    """
    Maakt een dictionary met als sleutels alle knopen van G en als value de degree van de sleutel in G
    """
    return {node: sum(1 for neighbor in G.nodes() if G.has_edge(node, neighbor)) for node in G.nodes()}


# test 1
print('G:', my_degree(G))

# test 2
ER = nx.gnm_random_graph(10, 20) # Zie https://networkx.github.io/documentation/stable/auto_examples/graph/plot_erdos_renyi.html
nx.draw(ER, with_labels=True);
my_degree(ER)
```

<!-- grade_id: Vraag4_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(my_degree(G), dict(G.degree()))
```

## Vraag 5
Teken de subgraven van `G` geinduceerd door {1,2,3} en {2,5,3} en {5,4} en geef deze graven de namen respectievelijk `I`, `J` en `K`

<!-- grade_id: Vraag5, points: 0, grade: false, solution: true -->
```python
# Subgraaf met knopen {1, 2, 3}
I = G.subgraph({1, 2, 3})
# Subgraaf met knopen {2, 5, 3}
J = G.subgraph({2, 5, 3})
# Subgraaf met knopen {5, 4}
K = G.subgraph({5, 4})

#WRITE YOUR CODE HERE
plt.figure(1)
nx.draw(I)
plt.figure(2)
nx.draw(J)
plt.figure(3)
nx.draw(K)
plt.show()
```

<!-- grade_id: Vraag5_antwoord, points: 1, grade: true, solution: false -->
```python
assert_equal(sorted(list(I.nodes())), [1, 2, 3])
assert_equal(sorted(list(J.nodes())), [2, 3, 5])
assert_equal(sorted(list(K.nodes())), [4, 5])
```

## Vraag 6
Maak de functie `adj_matrix(G)`  en `adj_matrix2(G)` die allebei een adjacency matrix maken. In het geval van de graaf `G` die je hebt gemaakt bij vraag 1, heeft de matrix van deze  vorm (alleen dan in de vorm van een lijst van lijsten):
    <pre>
        1 2 3 4 5
      1 0 1 0 1 0
      2
      3
      4
      5
    </pre>
    
Een aantal eisen:
* Zorg dat de functies werken voor elk netwerk, niet alleen `G`. (We kijken de functies ook na met andere netwerken dan `G`.
* Gebruik enkel list comprehensies.
* Je mag in beide functies niet de networkx functie `nx.adjacency_matrix` gebruiken.
* Bij `adj_matrix(G)` moet je de networkx functie `G.has_edge` gebruiken en  geen enkele andere networkx functie.
* Bij `adj_matrix2(G)` moet je de networkx functie `G.neighbors` (dit is niet hetzelfde als `nx.neighbors`!) gebruiken en  geen enkele andere networkx functie.
* Zorg dat de nodes altijd gesorteerd zijn, zodat de matrix op volgorde is
* Hint:`int(True)==1` en `int(False)==0`, en dus ook `int(5==4)==0`.
    
Je antwoord voor netwerk `G` zal er zo uit zien:

```
[[0, 1, 0, 1, 0],
 [....],
       ....
 [....]
]
```

<!-- grade_id: Vraag6, points: 0, grade: false, solution: true -->
```python
def adj_matrix(G):
    # Adjacency matrix algorithm implemented with G.has_edge

    nodes = G.nodes()

    return [[
            int(G.has_edge(y_node, x_node))
            for x_node in nodes
        ] for y_node in nodes
    ]


def adj_matrix2(G):
    # Adjacency matrix algorithm implemented with G.neighbours

    nodes = G.nodes()

    return [[
            int(x_node in G.neighbors(y_node))
            for x_node in nodes
        ] for y_node in nodes
    ]

adj_matrix(G),adj_matrix2(G), adj_matrix(G)==adj_matrix2(G)

# test het zelf ook met de ER graaf.
```

<!-- grade_id: Vraag6_test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(adj_matrix(G)), list)
assert_equal(type(adj_matrix(G)[0]), list)
assert_equal(type(adj_matrix(G)[0][0]), int)
assert_equal(type(adj_matrix2(G)), list)
assert_equal(type(adj_matrix2(G)[0]), list)
assert_equal(type(adj_matrix2(G)[0][0]), int)
assert_equal(type(adj_matrix2(G)[0][0]), int)
assert_array_almost_equal(adj_matrix(G), adj_matrix2(G))

```

# Opgaven uit First Course

1. Alles uit Hfd 0
2. Hfd 1: alles tot en met 21.
3. Maak opgave 21 in `networkx`
        * maak het netwerk, met `G.add_edges`
        * print het
        * maak ook de adjecency matrix als een lijst van lijsten
        * check dat deze adjencency matrix hetzelfde is als een adjececy matrix die networkx maakt.
4. Schrijf nu ook checks voor alle vragen 1.17-1.20 in Python en controleer je antwoord.

<!-- grade_id: FCa, points: 1, grade: true, solution: true -->
# Chapter 0

1. (b) Road Segments
2. (b) Street Intersections
3. A flight (route) from one airport to another.
4. (d) Hub Nodes with many links.
5. Undirected
6. Directed

# Chapter 1

1. Done
2. Given a single edge, the maximum amount of nodes that can be connected is N=2. The maximum amount of edges a single node can have, is N-1 where N is the amount of nodes in the graph.
3. The most common degree is `k = 4`.
4. The most common in-degree is `k = 2`, the most common out-degree is `k = 2`.
5. We can assign a weight to each of the links, where the higher the weight, the higher traffic.
6. (c) Total in-degree must be equal to the out-degree.
7. (d) Directed, Weighted
8. (b) Undirected, Weighted
9. (a) Undirected, Unweighted
10. (d) Directed, Weighted
11. Such a network is called a clique (1.4 "subnetworks"), the total amoun of links is `(25 * 24) / 2 = 300` links.
12. The maximum number of links an undirected network of `N` nodes can have at most `Lmax = N * (N-1) / 2`.
13. Given an amount of nodes `N`, divided into groups `Na` and `Nb` the maximum amount of nodes possible in a bipartite network is `Na * Nb`, since any nodes in group `a` can only have as many links as there are nodes in group `b`.
14. (a) Network A has more links than network B.
15. Yes, if there are less than `N*(N-1)/2` links in a graph there will be nodes that are not connected.
16. (d) Directed, Weighted
17. There a `N=6` nodes, with `L=10` links. There are no self-loops (links that originate from and point to the same node).
18. Node `F` has out-going links to every other node. There are no nodes with in-going links from all other nodes.
19. Only node `C` is a sink.
20. `S in(C) = 5`, and `S out(C) = 0`
21. There are `N=6` nodes, and `L=9` links, code is below:

```python
# Chapter 1, Exercise 21
adj_graph = nx.Graph()

# Adding the edges of the new undirected, unweighted graph
adj_graph.add_edge("A", "B")
adj_graph.add_edge("B", "C")
adj_graph.add_edge("D", "B")
adj_graph.add_edge("D", "E")
adj_graph.add_edge("E", "F")
adj_graph.add_edge("F", "A")
adj_graph.add_edge("F", "B")
adj_graph.add_edge("F", "C")
adj_graph.add_edge("F", "D")

# Drawing the graph.
nx.draw(adj_graph)

# Adjacency matrix made from a list of lists.
adj_matrix = [
    [0, 1, 0, 0, 0, 1],
    [1, 0, 1, 1, 0, 1],
    [0, 1, 0, 0, 0, 1],
    [0, 1, 0, 0, 1, 1],
    [0, 0, 0, 1, 0, 1],
    [1, 1, 1, 1, 1, 0],
]

# Adjacency matrix data from the networkx graph.
nx_adj_matrix = nx.adjacency_matrix(adj_graph)

# Check whether the matrices are the same.
print("Adjacency matrices are the same: ",
  (nx_adj_matrix.todense() == adj_matrix).all())
```

```python
# Checks for exercises 1.17-1.20

# Create the directional graph with weighted edges.
di_graph = nx.DiGraph()

di_graph.add_edges_from([
    ("A", "B", {"weight": 1}),
    ("B", "C", {"weight": 2}),
    ("D", "B", {"weight": 1}),
    ("D", "E", {"weight": 1}),
    ("E", "F", {"weight": 1}),
    ("F", "A", {"weight": 2}),
    ("F", "B", {"weight": 1}),
    ("F", "C", {"weight": 3}),
    ("F", "D", {"weight": 1}),
    ("F", "E", {"weight": 1}),
])

# Checks 1.17
print("Checks for 1.17:")
print("The amount of nodes should be 6: ", len(di_graph.nodes()) == 6)
print("The amount of links (edges) should be 10: ", len(di_graph.edges()) == 10)
print("There are no self-loops: ", not any(a == b for a, b in di_graph.edges()))

# Checks 1.18
print("\nChecks for 1.18:")
print("Only node F has out-going links to all other nodes: ",
      [
          node_A
          for node_A in di_graph.nodes() if all(
              di_graph.has_edge(node_A, node_B)
              for node_B in di_graph.nodes()
              if node_B != node_A
          )
      ] == ["F"])
print("No nodes have in-going links from all other nodes: ",
      len([
          node_A
          for node_A in di_graph.nodes() if all(
              di_graph.has_edge(node_B, node_A)
              for node_B in di_graph.nodes()
              if node_B != node_A
          )
      ]) == 0)

# Checks 1.19
print("\nChecks for 1.19:")
print("Only node C is a sink: ",
      [
          node
          for node in di_graph.nodes()
          if len(list(di_graph.neighbors(node))) == 0
      ] == ["C"])

# Checks 1.20
print("\nChecks for 1.20:")
print("The in-strength of node C is 5: ",
      sum(
          di_graph[node]["C"]["weight"]
          for node in di_graph.nodes()
          if di_graph.has_edge(node, "C")
      ) == 5)
print("The out-strength of node C is 0: ",
      len(list(di_graph.neighbors("C"))) == 0,
      )

```

# FC vraag 1.13

* Maak deze vraag ook in networkx.
    * Maak een functie `Q113(N1,N2)` die een network teruggeeft zoals gevraagd
    * Met welke check check je dat het "volledig" (voor een bipartiet netwerk natuurlijk) is?
    
* **Hint** Google eens wat.

<!-- grade_id: fc113a, points: 0, grade: false, solution: true -->
```python
from itertools import product
from networkx.algorithms import bipartite

def Q113(N1, N2):
    '''Return a maximal bipartite network with N1 nodes on one side and N2 on the other.'''
    return bipartite.complete_bipartite_graph(N1, N2)

def maximaal_bipartiet(BG):
    '''Return True if BG is a maximal bipartite network. False in all other cases.'''
    if not bipartite.is_bipartite(BG):
      return False

    node_sets = nx.bipartite.sets(BG)

    for a in node_sets[0]:
      for w in node_sets[1]:
        if not nx.has_path(BG, a, w):
          return False

    return True

grap = Q113(5, 10)
maximaal_bipartiet(Q113(5,10))
```

<!-- grade_id: fc113t, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q113(5,10),nx.Graph)
nx.draw(Q113(5, 10))
assert isinstance(maximaal_bipartiet(Q113(5,10)), bool)
```

<!-- grade_id: adjmatrixchecksd, points: 1, grade: true, solution: true -->
```python
#WRITE YOUR CODE HERE
```

<!-- grade_id: FCa21, points: 1, grade: true, solution: true -->
```python
#WRITE YOUR CODE HERE
```

# Verzamelingen


* Dubbel klik eens op deze cel zodat je de code ziet.
* Leer dan de latex codes voor de mooie symbolen, en zorg dat je zelf nu ook netjes antwoord kunt geven.
* Formules zet je altijd tussen dollartekens.
    * Als je het in het midden van de regel wil tussen dubbele dolalartekens
* Elk commando begint altijd met `\`. Spaties maken niks uit. Latex regelt hoe het er netjes uit komt te zien.

### Opgaven

1. $\{1,2,3\}\cap\{2,3,4\}=${2, 3}
1. $\{1,2,3\}\cup\{2,3,4\}=${1, 2, 3, 4}
1. $\{1,2,3\}\setminus\{2,3,4\}=${1}
1. $\emptyset\setminus\{2,3,4\}=$∅
2. $\{1,2,3\}\times\{2,3,4\}=${(1, 2), (1, 3), (1, 4), (2, 2), (2, 3), (2, 4), (3, 2), (3, 3), (3, 4)}
3. De machtsverzameling van $\{1,2,3,4\}$ is
{{}, {1}, {2}, {3}, {4}, {1, 2}, {1, 3}, {1, 4}, {2, 3}, {2, 4}, {3, 4}, {1, 2, 3}, {1, 2, 4}, {1, 3, 4}, {2, 3, 4}, {1, 2, 3, 4}}

#### Teken ook netjes de Venn diagrammen
#### Check je antwoorden ook in Python
#### Maak zelf opgaven en leg die aan je groepsleden voor

# Comprehensies

* Zorg dat je list, set en dict comprehensies echt goed in je vingers krijgt!
* Het is heerlijk als je het kunt.

### Goede uitleg

Ik [Googlede op python comprehension tutorial](https://www.google.com/search?q=python+comprehension+tutorial) en de [video](https://www.youtube.com/watch?v=3dt4OGnU5sM) is prima, maar ook de [datacamp tutorial](https://www.datacamp.com/community/tutorials/python-list-comprehension) of de [real python](https://realpython.com/list-comprehension-python/) tutorial.

### Wat opgaven

Je gaat allemaal comprehensies doen op een gedicht van annie MG Schmidt. Het staat in de cel hieronder als string in de variabele `AMGS` en daarna als lijst van woorden in `t`.

In je tests gebruik je allemaal string functies, en ook `in`.

#### list comprehensies

1. de lijst van alle woorden in `AMGS` beginnend met een `i`. (antwoord beneden)
2. zelfde maar dan beginnend met `i` of `I`. (antwoord beneden)
3. de lijst van alle woorden in `AMGS` van lengte meer dan 3
4. de lijst van alle woorden in `AMGS` waar een `a` in zit
5. nu waar juist geen `a` in zit
6. nu waar als er een `a` in zit ook een `b` in zit (wat betekent dit precies he?)
7. alle woorden zonder klinker.
8. Voor elk woord zijn lengte.
9. De lijst met het hele gedicht maar nu waar elk woord met een hoofdletter begint.
10. De lijst die net zo lang is als het gedicht (dus `t`), maar dan op elke positie of het woord op die positie gelijk is aan `ik`. Dus de waardes zijn alleen maar `True` en `False`.
    * Gebruik `sum` om te berekenen hoe vaak het woord `ik` voorkomt.
11. Verzien er zelf 10 bij, en maak die ook, en geef ze aan je groepsgenoeten.

#### set comprehensies

1. De verzameling van alle unieke woorden in `AMGS`.
2. De verzameling van alle unieke woorden in `AMGS` met `ch` er in.
3. De verzameling van alle unieke woorden in `AMGS`, maar dan geschreven in hoofdletters.
4. Alle woordlengtes die in `AMGS` voorkomen.


### Dict comprehensies

1. De dict met voor elk uniek "woord" in `AMGS` zijn lengte. Dus het woord is de key en de lengte de waarde. Let op , woorden zijn dus de elementen in `t`. (antwoord beneden)
2. De dict met voor elk uniek "woord" in `AMGS` de verzameling karakters in dat woord.
3. De dict met voor elk uniek "woord" in `AMGS`  de "schoongemaakte" versie van dat woord (zonder leestekens, etc, gebruik een lange lijst als `a.replace(':','').replace(',','')`. Zet het woord ook in klein eletters.

    
### Nesting

3. De dict met voor elk uniek woord, de verzameling klinkers in dat  woord.
    *   (antwoord beneden, probeer eerst minstens een half uur zelf)
4.  De dict met voor elke positie in `AMGS`, de verzameling met het woord voor die positie  en het woord erna.
    (antwoord ook beneden, niet meteen gan kijken.)
5. DE dict met voor elk woord de verzameling posities waar het voorkomt. Tricky maar leuk!   Hint loop over de posities.
6. Na de vorige is deze heel makkelijk: DE dict met voor elk woord  hoe vaak het voorkomt.


```python
AMGS='''
Ik wil niet meer, ik wil niet meer!
Ik wil geen handjes geven!
Ik wil niet zeggen elke keer:
Jawel mevrouw, jawel meneer...
nee, nooit meer in m'n leven!
Ik hou m'n handen op m'n rug
en ik zeg lekker niks terug!

Ik wil geen vieze havermout,
ik wil geen tandjes poetsen!
Ik wil lekker knoeien met het zout,
ik wil niet aardig zijn, maar stout
en van de leuning roetsen
en schipbreuk spelen in de teil
en ik wil spugen op het zeil!

En heel hard stampen in een plas
en dan m'n tong uitsteken
en morsen op m'n nieuwe jas
en ik wil overmorgen pas
weer met twee woorden spreken!
En ik wil alles wat niet mag,
de hele dag, de hele dag!

En ik wil op de kanapee
met hele vuile schoenen
en ik wil aldoor gillen: nee!
En ik wil met de melkboer mee
en dan het paardje zoenen.
En dat is alles wat ik wil
en als ze kwaad zijn, zeg ik: Bil!'''

t=AMGS.split()
len(t),t[:10]
```

```python
####list comprehensions

#1. de lijst van alle woorden in `AMGS` beginnend met een `i`. (antwoord beneden)
#2. zelfde maar dan beginnend met `i` of `I`. (antwoord beneden)

#3. de lijst van alle woorden in `AMGS` van lengte meer dan 3
#[w for w in t if len(w) >3]

#4. de lijst van alle woorden in `AMGS` waar een `a` in zit
#[w for w in t if 'a' in w]

#5. nu waar juist geen `a` in zit
#[w for w in t if 'a' not in w]

#6. nu waar als er een `a` in zit ook een `b` in zit (wat betekent dit precies he?)
#[w for w in t if 'a' in w and 'b' in w]

#7. alle woorden zonder klinker.
#[w for w in t if all(l not in 'aeiouAEIOU' for l in w)]

#8. Voor elk woord zijn lengte.
#[len(w) for w in t]

#9. De lijst met het hele gedicht maar nu waar elk woord met een hoofdletter begint.
#[w[0].upper() + w[1:] for w in t]

#10. De lijst die net zo lang is als het gedicht (dus `t`), maar dan op elke positie of het woord op die positie gelijk is aan `ik`. Dus de waardes zijn alleen maar `True` en `False`.
#   * Gebruik `sum` om te berekenen hoe vaak het woord `ik` voorkomt.
#sum([w == 'ik' for w in t ])

#11. Verzien er zelf 10 bij, en maak die ook, en geef ze aan je groepsgenoeten.
#:'(
```

```python
#### set comprehensies

#1. De verzameling van alle unieke woorden in `AMGS`.
#{w for w in t}

#2. De verzameling van alle unieke woorden in `AMGS` met `ch` er in.
#{w for w in t if 'ch' in w}

#3. De verzameling van alle unieke woorden in `AMGS`, maar dan geschreven in hoofdletters.
#{w.upper() for w in t}

#4. Alle woordlengtes die in `AMGS` voorkomen.
#{len(w) for w in t}
```

```python
### Dict comprehensies

#1. De dict met voor elk uniek "woord" in `AMGS` zijn lengte. Dus het woord is de key en de lengte de waarde. Let op , woorden zijn dus de elementen in `t`. (antwoord beneden)
#{w: len(w) for w in t}

#2. De dict met voor elk uniek "woord" in `AMGS` de verzameling karakters in dat woord.
#{w: set(w) for w in t}

#3. De dict met voor elk uniek "woord" in `AMGS`  de "schoongemaakte" versie van dat woord (zonder leestekens, etc, gebruik een lange lijst als `a.replace(':','').replace(',','')`. Zet het woord ook in klein eletters.
#{w: w.lower().replace(',','').replace('!','') for w in t}
```

```python
### Nesting

#3. De dict met voor elk uniek woord, de verzameling klinkers in dat  woord.
#    *   (antwoord beneden, probeer eerst minstens een half uur zelf)
#{w:{l for l in w if l in 'aeiouAEIOU'} for w in set(t)}

#4.  De dict met voor elke positie in `AMGS`, de verzameling met het woord voor die positie  en het woord erna.
#    (antwoord ook beneden, niet meteen gan kijken.)
#{i:{t[i-1], t[i+1]} for i in range(1,len(t)-1)}

#5. DE dict met voor elk woord de verzameling posities waar het voorkomt. Tricky maar leuk!   Hint loop over de posities.
#{word: [i for i, w in enumerate(t) if w == word] for word in set(t)}

#6. Na de vorige is deze heel makkelijk: DE dict met voor elk woord  hoe vaak het voorkomt.
#{word: len([i for i, w in enumerate(t) if w == word]) for word in set(t)}
```

```python
# probeer eens alleen t[0]. en dan shift return, en probeer eens wat van de suggesties
# om erachter te komen wat iets doet typ je t[0].isupper? en dan shift enter.
# Ideaal he! Dat kan bij elk object. t[0] is natuurlijk een string, dus je krijgt string functies te zien.
t[0].isupper()
```

```python
# "in" is ook heel handig in Python

print(t[1])
'i' in t[1], 'W' in t[1], 'W' in t[1].upper()
```

```python
#de lijst van alle woorden in `AMGS` beginnend met een `i`. (antwoord beneden)
[w for w in t if w.startswith('i')]
```

```python
#2. zelfde maar dan beginnend met `i` of `I`. (antwoord beneden)
# Verzin minstens 2 andere manieren om exact hetzelfde te krijgen
[w for w in t if w.startswith('i') or w.startswith('I')]
```

```python
#  De dict met voor elk uniek woord in `AMGS` zijn lengte.

{w:len(w) for w in t}
```

```python
# De dict met voor elk uniek woord, de verzameling klinkers in dat  woord.

{w:{l for l in w if l in 'aeiou'}
    for w in set(t)} # verbeter hem zelf zodat hoofdletters ook mee worden genomen

```

```python
#  De dict met voor elke positie in `AMGS`, de verzameling met het woord voor die positie  en het woord erna.

{i:{t[i-1],t[i+1]} for i in range(1,len(t)-1)}
```

<!-- grade_id: TRICKY, points: 0, grade: false, solution: true -->
```python
#  DE dict met voor elk woord de verzameling posities waar het voorkomt. Tricky maar leuk!
l=dict()
#WRITE YOUR CODE HERE
```

```python


```